# HW1: Ẩn tin mật trên văn bản bằng phương pháp chèn khoảng trắng <font color="grey">(chèn khoảng trắng giữa các từ để làm cho văn bản được căn lề 2 bên)

(Cập nhật lần cuối: 24/02/2023)

Họ tên: Nguyễn Quốc Huy

MSSV: 20127188

---

## Nắm rõ cách làm bài và nộp bài
&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để phần nào giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.

*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>.  Bạn có thể thảo luận ý tưởng với bạn khác cũng như tham khảo các nguồn trên mạng, nhưng sau cùng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font> (khi tham khảo các nguồn trên mạng thì bạn cần ghi rõ nguồn trong bài làm, và đương nhiên là bạn cũng không được phép đưa code và bài làm cho bạn khác xem). <font color=red>Nếu vi phạm những điều này thì có thể bạn sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `HW1.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại với định dạng nén là .zip (chứ không được là .rar hay các định dạng khác) và nộp ở link trên moodle. \
<font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài này.</font>

## Kiểm tra môi trường code

```python
import sys
sys.executable
```

Output:
```
'c:\\Users\\huyng\\.conda\\envs\\dh_ss-env\\python.exe'
```

Bạn nên thấy kết quả in ra là đường dẫn đến file chạy python của môi trường "dh_ss-env" mà mình đã hướng dẫn bạn cài đặt ở HW1.

## Import các thư viện cần thiết

```python

```

## Cài đặt các hàm hỗ trợ

```python
def convert_s2bits(s):
    '''
    Chuyển chuỗi ký tự (ascii) thành chuỗi bit.
    '''
    bits = ''
    for c in s:
        c_bits = bin(ord(c))[2:].zfill(8)
        bits = bits + c_bits
    return bits

def convert_bits2s(bits):
    '''
    Chuyển chuỗi bit thành chuỗi ký tự (ascii).
    '''
    s = ''
    for i in range(0, len(bits), 8):
        c = chr(int(bits[i:i+8], 2))
        s = s + c
    return s
```

```python
# YOUR CODE HERE (OPTION)
```

## Cài đặt hàm nhúng (6đ)

```python
def embed(msg_file, cover_text_file, text_width, stego_text_file):
    '''
    Nhúng tin mật vào văn bản bằng phương pháp chèn khoảng trắng để văn bản 
    được căn lề 2 bên (mã giả ở trang 17 của file slide lý thuyết "02..."). 
    Ghi chú: trong trường hợp m >= n (xem trang 17 của file slide lý thuyết "02...") 
    thì bạn chèn khoảng trắng theo qui ước: duyệt dòng từ đầu về cuối, với 
    mỗi vị trí giữa các từ thì chèn thêm một khoảng trắng; lặp lại thao tác này 
    cho tới khi hết khoảng trắng cần chèn cho dòng. Ví dụ: với m = 12 và n = 10 
    thì 2 vị-trí-giữa-các-từ đầu sẽ được chèn 2 khoảng trắng, các 
    vị-trí-giữa-các-từ còn lại sẽ được chèn 1 khoảng trắng.
    
    Các tham số:
        msg_file (str): Tên file chứa secret message.
        cover_text_file (str): Tên file chứa cover text.
        text_width (int): Chiều dài của dòng sau khi được căn lề.
        stego_text_file (str): Tên file chứa stego text.
    Giá trị trả về:
        bool: True nếu nhúng thành không, False nếu không đủ chỗ để nhúng. 
    '''
    # Đọc file gốc
    with open(cover_text_file, 'r') as f:
        cover_text_lines = f.read().splitlines()
    
    # Đọc file thông điệp
    with open(msg_file, 'r') as f:
        msg = f.read()
    
    # Đổi thông điệp về dạng bit
    msg_bits = convert_s2bits(msg)

    # Khởi tạo
    b = 0 
    l = 0
    flag = False
    stego_text = ""
    
    # Chép thông tin vào biến để xử lý
    while l < len(cover_text_lines):
        line = cover_text_lines[l].strip()

        if l == len(cover_text_lines) - 1: 
            stego_text += line
            break
        elif cover_text_lines[l + 1] == "":
            stego_text += (line + "\n")
            l += 1
            continue

        # Tách từ và câu 
        words = line.split(' ')
        n = line.count(' ')
        m = text_width - len(line)
        num_spaces_inserted = 0

        # Tính lượng bit cần nhúng trong câu
        if min(m, n - m) > 0:
            bit_count_to_embed = min(m, n - m)
        else:
            bit_count_to_embed = 0
        next_stego_line = ""
        remaining_word = len(words)
        remaining_space = n + m
        
        # Kiểm tra điều kiện câu
        if 0 < m < n:
            for i in range(bit_count_to_embed):
                if b < len(msg_bits):
                    # Thực hiện nhúng từng bit
                    if msg_bits[b] == '1':
                        next_stego_line += (words[i * 2] + " " + words[i * 2 + 1] + "  ")
                    else:
                        next_stego_line += (words[i * 2] + "  " + words[i * 2 + 1] + " ")
                    b += 1
                else:
                    if flag:
                        next_stego_line += (words[i * 2] + "  " + words[i * 2 + 1] + " ")
                    else:
                        next_stego_line += (words[i * 2] + " " + words[i * 2 + 1] + "  ")
                        flag = True
                num_spaces_inserted += 3
            
            remaining_word -= bit_count_to_embed * 2
            remaining_space -= num_spaces_inserted
        
        if(remaining_word > 1):
            for i in range(remaining_word):
                num_space = 0
                if(i != remaining_word - 1):
                    num_space_each = int(remaining_space / (remaining_word - 1))
                    if (i < (remaining_space - (remaining_word - 1) * num_space_each)):
                        num_space_addition = 1
                    else:
                        num_space_addition = 0
                    num_space = num_space_each + num_space_addition
                next_stego_line += (words[bit_count_to_embed * 2 + i] + " " * num_space) 
        else:
            next_stego_line += words[bit_count_to_embed * 2]
            
        stego_text += (next_stego_line + "\n")
        l += 1

    if not flag:
        return False
        
    # Ghi thông tin đã nhúng vào file
    with open(stego_text_file, 'w') as f:
        f.write(stego_text)   
    
    return True

```

<!-- grade_id: cell-057a58bf038e1e60, points: 5, grade: true, solution: false -->
```python
# TEST 1            
result = embed('msg.txt', 'cover.txt', 70, 'stego.txt')
assert result == True
with open('stego.txt', 'r') as f:
    stego_text = f.read()
with open('correct_stego.txt', 'r') as f:
    correct_stego_text = f.read()
assert stego_text == correct_stego_text
```

<!-- grade_id: cell-f625774497df6542, points: 1, grade: true, solution: false -->
```python
# TEST 2
result = embed('msg2.txt', 'cover.txt', 70, 'stego.txt')
assert result == False
```

## Cài đặt hàm rút trích (4đ)

<!-- grade_id: cell-f3f9b34770f282aa, points: 0, grade: false, solution: true -->
```python
def extract(stego_text_file, extr_msg_file):
    '''
    Extracts a secret message from stego text without using bitarray.

    Parameters:
        stego_text_file (str): File name containing the stego text.
        extr_msg_file (str): File name to store the extracted secret message.

    Returns:
        None
    '''
    # Đọc file văn bản stego
    with open(stego_text_file, 'r') as f:
        stego_text_lines = f.read().splitlines()
    
    # Trích xuất các bit thông điệp
    message_bits = []
    for line in stego_text_lines:
        words = line.split(' ')
        if len(words) > 2:
            spaces = []
            for i in range(len(words) - 1):
                if words[i] != '':
                    if words[i + 1] == '':
                        spaces.append(2)
                    else:
                        spaces.append(1)
            for j in range(0, len(spaces), 2):
                if len(spaces[j : j + 2]) == 2:
                    if spaces[j] == 2 and spaces[j + 1] == 1:
                        message_bits.append(0)
                    elif spaces[j] == 1 and spaces[j + 1] == 2:
                        message_bits.append(1)

    # Tìm vị trí kết thúc của thông điệp
    last_index = len(message_bits) - 1
    for i in reversed(range(len(message_bits))):
        if message_bits[i] == 1:
            last_index = i
            break
    
    # Chuyển các bit thành thông điệp
    extracted_bits = ''.join(map(str, message_bits[:last_index]))
    extracted_message = convert_bits2s(extracted_bits)
    
    # Ghi thông điệp trích xuất vào file
    with open(extr_msg_file, 'w') as f:
        f.write(extracted_message)
```

<!-- grade_id: cell-caf1afe26d95b670, points: 4, grade: true, solution: false -->
```python
# TEST
extract('correct_stego.txt', 'extr_msg.txt')
with open('extr_msg.txt', 'r') as f:
    extr_msg = f.read()
with open('correct_extr_msg.txt', 'r') as f:
    correct_extr_msg = f.read()
assert extr_msg == correct_extr_msg
```